# Assignment 0

This is just a test assignment. Don't worry, it's not for any points! We're just making sure everything works.

There's one question below--a small code snippet. Don't worry about what it does; instead, focus on the double-quotes. See how they're right next to each other? Put something between those quotes. Anything at all. The autograder for this question will literally look to see if there's so much as a single blank space between the quotes, and will call the answer correct. `" "` would be correct; so would `"this is a lot more than just a space"`. Make something up! But whatever you do, *don't delete the quotes*.

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
def things_and_stuff(some_string):
    return "Here's what you wrote: {}".format(some_string)

instring = "this passes!"
# 1: Copy the *whole* line above this one (the one that says "instring")
# 2: Paste it on top of the "raise NotImplementedError" (as in, paste over it)
# 3: Put something in the double-quotes! Anything!
### BEGIN SOLUTION

### END SOLUTION
out = things_and_stuff(instring)
```

<!-- grade_id: q1_check, points: 1, grade: true, solution: false -->
```python
assert type(out) == str
assert len(out) > 0
```